### Part 1: Reading Data

```python
import hw3
```

```python
from hw3 import read_movies_data, read_ratings_data
```

```python
movies_df = read_movies_data("moviesSample.txt")
```

<!-- grade_id: test_read_movies_data, points: 8, grade: true, solution: false -->
```python
assert movies_df.shape == (9, 3)
assert movies_df.iloc[2, 2] == "Comedy"
assert movies_df.loc[6, "title"] == "Heat"
assert all(movies_df.columns == ["title", "year", "genre"])
assert list(movies_df.index) == [1, 2, 3, 4, 5, 6, 8, 9, 10]
```

```python
movie_ratings = read_ratings_data("ratingsSample.csv")
```

<!-- grade_id: test_read_ratings_data, points: 8, grade: true, solution: false -->
```python
assert sum(movie_ratings[5]) == 24
```

### Part 2: Processing Data

```python
from hw3 import get_movie, create_genre_dict, calculate_average_rating
```

```python
genre_dict = create_genre_dict(movies_df)
```

```python
avg_ratings = calculate_average_rating(movie_ratings, movies_df)
```

<!-- grade_id: test_calculate_average_rating, points: 8, grade: true, solution: false -->
```python
assert avg_ratings.loc[3] == 4.0
```

### Part 3: Recommendations

```python
from hw3 import get_popular_movies, filter_movies, get_popular_in_genre
from hw3 import get_genre_rating, get_movie_of_the_year
```

```python
get_popular_movies(avg_ratings, 4)
```

Output:
```
6    4.250000
5    4.000000
3    4.000000
1    3.833333
dtype: float64
```

<!-- grade_id: test_get_popular_movies, points: 8, grade: true, solution: false -->
```python
pop_movies = get_popular_movies(avg_ratings, 4)
assert pop_movies.is_monotonic_decreasing
assert set(pop_movies.index) == {6, 5, 3, 1}
```

```python
filter_movies(avg_ratings, 3.2)
```

Output:
```
{1: 3.8333333333333335,
 2: 3.4166666666666665,
 3: 4.0,
 5: 4.0,
 6: 4.25,
 9: 3.3333333333333335}
```

<!-- grade_id: test_filter_movies, points: 8, grade: true, solution: false -->
```python
good_movies = filter_movies(avg_ratings, 3.2)
# assert good_movies.shape[0] == 6
assert len(good_movies) == 6 #changed this since shape did not work for dict
# assert all(good_movies.values >= 3.2)
assert all(value >= 3.2 for value in good_movies.values()) #changed this since value is not an attribute
```

```python
get_popular_in_genre("Adventure", genre_dict, avg_ratings, 2)
```

Output:
```
1    3.833333
2    3.416667
dtype: float64
```

<!-- grade_id: test_get_popular_in_genre, points: 8, grade: true, solution: false -->
```python
pop_in_genre = get_popular_in_genre("Adventure", genre_dict, avg_ratings, 2)
assert list(pop_in_genre.index) == [1, 2]
assert pop_in_genre.is_monotonic_decreasing
```

```python
get_genre_rating("Action", genre_dict, avg_ratings)
```

Output:
```
np.float64(3.527777777777778)
```

```python
genre_rating = get_genre_rating("Action", genre_dict, avg_ratings)
assert genre_rating > 3.525
assert genre_rating < 3.53
```

<!-- grade_id: test_get_movie_of_the_year, points: 8, grade: true, solution: false -->
```python
assert get_movie_of_the_year(1995, avg_ratings, movies_df) == 'Heat'
```

```python
from hw3 import read_user_ratings, get_user_genre, recommend_movies
```

```python
user_movies = read_user_ratings("ratingsSample.csv")
```

<!-- grade_id: test_read_user_ratings, points: 8, grade: true, solution: false -->
```python
assert user_movies[42] == [(3, 4.0)]
```

Function `get_user_genre` that takes as parameters a user id, the user-to-movies dictionary
(as created in Part 4.1 above), and the movies DataFrame (as created in Part 1.1), and returns
the top genre that the user likes based on the user’s ratings. Here, the top genre for the user will be
determined by taking the average rating of the movies genre-wise that the user has rated. If multiple
genres have the same highest ratings for the user, return any one of genres (arbitrarily) as the top
genre.


<!-- grade_id: test_get_user_genre, points: 10, grade: true, solution: false -->
```python
user_genre = get_user_genre(42, user_movies, movies_df)
assert user_genre == "Comedy"
```

Recommend 3 most popular (highest average rating) movies from the user’s top genre that the user
has not yet rated. Write a function called `recommend_movies` for this, that takes as parameters a user id,
the user-to-movies dictionary (as created in Part 4.1 above), the movies DataFrame (as created
in Part 1.1), and the movie-to-average rating Series (as created in Part 2.2). The function should
return a pandas Series of movie-to-average rating, sorted in decreasing order by rating. If fewer than 3 movies make the cut, then return all
the movies that make the cut in order of top average ratings.

```python
recommend_movies(6, user_movies, movies_df, avg_ratings)
```

Output:
```
Series([], dtype: float64)
```